# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 2: Clustering

```ruby
require './assignment_lib'
```

Output:
```
"if(window['d3'] === undefined ||\n   window['Nyaplot'] === undefined){\n    var path = {\"d3\":\"https://cdnjs.cloudflare.com/ajax/libs/d3/3.5.5/d3.min\",\"downloadable\":\"http://cdn.rawgit.com/domitry/d3-downloadable/master/d3-downloadable\"};\n\n\n\n    var shim = {\"d3\":{\"exports\":\"d3\"},\"downloadable\":{\"exports\":\"downloadable\"}};\n\n    require.config({paths: path, shim:shim});\n\n\nrequire(['d3'], function(d3){window['d3']=d3;console.log('finished loading d3');require(['downloadable'], function(downloadable){window['downloadable']=downloadable;console.log('finished loading downloadable');\n\n\tvar script = d3.select(\"head\")\n\t    .append(\"script\")\n\t    .attr(\"src\", \"http://cdn.rawgit.com/domitry/Nyaplotjs/master/release/nyaplot.js\")\n\t    .attr(\"async\", true);\n\n\tscript[0][0].onload = script[0][0].onreadystatechange = function(){\n\n\n\t    var event = document.createEvent(\"HTMLEvents\");\n\t    event.initEvent(\"load_nyaplot\",false,false);\n\t    win
```

## Question 1.1: K-Means (25 points)

The following dataset has 5 clusters. Generate a dataset with 5 means in two dimensions, with means $\mu$ and variance $\sigma$ defined as follows:
* $\mu_1 = (-4,4)$, $\sigma_1 = 1$
* $\mu_2 = (-4,-4)$, $\sigma_2 = 1$
* $\mu_3 = (4,4)$, $\sigma_3 = 1$
* $\mu_4 = (4,-4)$, $\sigma_4 = 1$
* $\mu_5 = (0,0)$, $\sigma_5 = 1$

Name your features ```x1``` and ```x2```.

Generate 100 points in each cluster and plot each cluster with a different color.  Use the same format we usually use with the cluster id as the ```label``` field and also set the cluster in a field called ```cluster```. 

To generate a value from a normal distribution, use the following:

```ruby
r = Distribution::Normal.rng(mean,stdev)
x = r.call()
```

<!-- grade_id: cell-1cb0100f7efe89e6, points: 0, grade: false, solution: true -->
```ruby
def create_cluster_dataset()
  n = 100
  clusters = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  dataset = []
  stdev = 1
  
  clusters.each.with_index do |mean, i|
    r = mean.collect{|m| Distribution::Normal.rng(m, stdev)}
    #iterate 100 times
     n.times do |j|
       dataset << {"features" => {"x1" => r[0].call(), "x2" => r[1].call()}, "label" => i, "cluster" => i}
     end
  end
  
#   dataset.size.times do |i|
#     puts dataset[i]
#   end
  
  return dataset
end
```

Output:
```
:create_cluster_dataset
```

<!-- grade_id: cell-1ac6e84306bf8a37, points: 5, grade: true, solution: false -->
```ruby
def test_11_1()
  dataset = create_cluster_dataset()
  assert_false(dataset.empty?)
  assert_equal 500, dataset.size
  plot_clusters(dataset)
end
test_11_1()
```

Output:
```
#<Nyaplot::Plot:0x00000000021a6218 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000001b68ce8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(179,226,205)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"ade9d23a-531f-421a-b73c-e6574b52e644"}, @xrange=[-7.189454285349745, -1.5110551633194769], @yrange=[0.8349409204528402, 6.561709157985823]>, #<Nyaplot::Diagram:0x0000000002223ad8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(253,205,172)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 
```

<!-- grade_id: cell-a6100841ca0c2971, points: 10, grade: true, solution: false -->
```ruby
def test_11_2()
  dataset = create_cluster_dataset()
  counts = dataset.group_by {|x| x["cluster"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}
  
  counts = dataset.group_by {|x| x["label"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}

end
test_11_2()
```

<!-- grade_id: cell-09a1705f91d1d892, points: 10, grade: true, solution: false -->
```ruby
def test_11_3()
  dataset = create_cluster_dataset()
  clusters = dataset.group_by {|x| x["label"]}
  means = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  5.times do |i|
    assert_not_nil clusters[i]
    assert_equal 100, clusters[i].size
    m1 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x1"]} / 100.0
    m2 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x2"]} / 100.0
    assert_in_delta m1, means[i][0], 0.5
    assert_in_delta m2, means[i][1], 0.5    
  end
end
test_11_3()
```

Output:
```
5
```

## Question 1.2 (25 points)

Generate $k = 5$ random means, each within the min and max of each feature value.

<!-- grade_id: cell-6a934e257380e926, points: 0, grade: false, solution: true -->
```ruby
def init_cluster data, k
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  rnd = Random.new(1234)

  k.times.with_index do |j|
    means[j]["x1"] = rnd.rand(-10.0..10.0) 
    means[j]["x2"] = rnd.rand(-10.0..10.0)
  end
  means
end
```

Output:
```
:init_cluster
```

<!-- grade_id: cell-52ceae8580209a6d, points: 25, grade: true, solution: false -->
```ruby
def test_12_1()
  dataset = create_cluster_dataset()  
  means = init_cluster dataset, 5
  assert_equal 5, means.size
  assert_equal 2, means.first.size
  5.times {|i| assert means.has_key? i}
  assert_true(means.values.all? {|m| m["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means.values.all? {|m| m["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_12_1()
```

## Question 2.1: (10 points)
We will now implement the k-means algorithm to recover the clusters above. Return the means and plot the obtained clustering. Compare your discovered means to the known means.

Starting from a randomly initialized set of clusters defined by a set of means, we will iteratively refine our estimate of the clustering. At each step, we update the means, assign clusters, and update means from the assigned clustering. We denote candidate cluster membership by the binary vector $z_{i,j} = 1$ if example $i$ belongs to cluster $j$. Use an ```Array``` for a row of z.

Start by setting the $z$ vector by assigning finding the mean closest to each point in the datset. Set the candidate cluster in the ```cluster``` field in each row. 

<!-- grade_id: cell-8703280af8c50289, points: 0, grade: false, solution: true -->
```ruby
def assign_cluster(data, means)
  indices = Array.new(means.size) {|i| i}
  z = Array.new
  group_num = means.size
  
  data.each.with_index do |record, i|
    min_mean = find_closest record["features"], means
    candi_cluster_arr = Array.new(group_num)
    candi_cluster_arr.fill(0.0)
    candi_cluster_arr[min_mean] = 1.0
    z << candi_cluster_arr
  end
  
#   z.size.times do |i|
#     puts z[i]
#   end
  return z
end


def find_closest(point, means)
  min = 0.0
  min_index = 0
  n = means.size
  
  n.times do |i|
    mean = means[i]
    curr_dist = calc_dist_1(point, mean)
    if(min == 0.0 || curr_dist < min)
      min = curr_dist
      min_index = i
    end
  end
  return min_index
end

def calc_dist_1(point1, point2)
  dist = Math.sqrt((point2["x1"] - point1["x1"]) ** 2 + (point2["x2"] - point1["x2"]) ** 2)
  return dist
end
```

Output:
```
:calc_dist_1
```

<!-- grade_id: cell-534b63511cbeda9d, points: 10, grade: true, solution: false -->
```ruby
def test_21_1()
  dataset = create_cluster_dataset()  
  means = init_cluster dataset, 5
  z = assign_cluster(dataset, means)
  assert_equal 500, z.size
  assert_true(z.all? {|zi| zi.size == 5}, "Must set a value for each cluster")
  assert_true(z.all? {|zi| g = zi.group_by {|zij| zij}; g[0.0].size == 4 and g[1.0].size == 1}, 
    "Must set only one cluster to 1.0 and all others to 0.0 (not an integer)")
  
  plot_clusters(dataset)
end
test_21_1()
```

Output:
```
#<Nyaplot::Plot:0x0000000002182228 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000012843c0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(141,211,199)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"00da9dac-cc6f-4b4c-8503-0110ff222379"}, @xrange=[-6.177682504367882, -2.1488124663192427], @yrange=[1.2070731638696088, 5.77297354193761]>, #<Nyaplot::Diagram:0x0000000001be48e8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(255,255,179)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1
```

## Question 2.2: (5 points)
Given the $z$ vector and the dataset, calculate the means for each cluster determined by $z$. 


<!-- grade_id: cell-f651c1c0bc2344d7, points: 0, grade: false, solution: true -->
```ruby
def calculate_means z, data
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  z.each.with_index do |candi_arr, index|
    candi_arr.each.with_index do |candi, i|
      if candi == 1.0
        data[index]["cluster"] = i
      end
    end
  end
  
  data_cluster = data.group_by{|x| x["cluster"]}
  
  #check if the cluster is empty
    data_cluster.size.times do |i|
#       puts "@#$%^&*()#$%^&*(#$%^&*())"
#       puts data_cluster[i].size
      if data_cluster[i] != nil
        means[i]["x1"] = data_cluster[i].inject(0.0){|avg, n| avg += n["features"]["x1"]} / data_cluster[i].size
        means[i]["x2"] = data_cluster[i].inject(0.0){|avg, n| avg += n["features"]["x2"]} / data_cluster[i].size
      end
    end
  
#    data_cluster.size.times do |i|
#      puts means[i]
#    end
  # BEGIN YOUR CODE
  # raise NotImplementedError.new()
  #END YOUR CODE
  means
end
```

Output:
```
:calculate_means
```

<!-- grade_id: cell-7fff6a6137b36c9b, points: 5, grade: true, solution: false -->
```ruby
def test_22_1()
  dataset = create_cluster_dataset()  
  means1 = init_cluster dataset, 5
  z = assign_cluster(dataset, means1)

  assert_equal 500, dataset.size
  assert_equal 500, z.size

  means2 = calculate_means(z, dataset)
  assert_equal 5, means2.size
  5.times {|i| assert means2.has_key? i}
  
  5.times do |i|
    assert_true((means2[i]["x1"] - means1[i]["x1"]).abs > 1e-5, "Expected means to change")
    assert_true((means2[i]["x2"] - means1[i]["x2"]).abs > 1e-5, "Expected means to change")
  end  

  assert_true(means2.values.all? {|fm| fm["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means2.values.all? {|fm| fm["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_22_1()
```

## Question 2.3: (5 points)
As we iteratively refine our means for the clusters, we need to terminate after some criterion. Stop when the cluster distances are less than $\tau = 0.001$ and plot the convergence. First, let's calculate the distance between the set of $k$ means, using squared distance. Use the following formula for the termination criteria:

#  $\frac{1}{k} \sum_{k} \left\lVert \mu^{(k)}_t - \mu^{(k)}_{t - 1} \right\rVert^2 \le \tau$


<!-- grade_id: cell-b7105f7e91a2eba1, points: 0, grade: false, solution: true -->
```ruby
def cluster_dist(m0, m1)
  dist = 0.0
  n = m0.size

  n.times do |i|
    dist_temp =  calc_dist m0[i], m1[i]
    dist += dist_temp / n
  end
  dist
end

def calc_dist(point1, point2)
  dist = (point2["x"] - point1["x"]) ** 2 + (point2["y"] - point1["y"]) ** 2
  return dist
end
```

Output:
```
:calc_dist
```

<!-- grade_id: cell-abe63804c50cf21b, points: 5, grade: true, solution: false -->
```ruby
def test_23_1()
  m0 = {0 => {"x" => 1.0, "y" => 2.0}, 1 => {"x" => 2.0, "y" => 3.0}}
  m1 = {0 => {"x" => 2.0, "y" => 2.0}, 1 => {"x" => 3.0, "y" => 4.0}}

  assert_equal 0.0, cluster_dist(m0, m0)
  assert_equal 0.0, cluster_dist(m1, m1)
  assert_equal 1.5, cluster_dist(m0, m1)
end

test_23_1()
```

## Question 2.4: (10 Points)

We are now ready to complete the $k$-means clustering algorithm. Given a dataset, the number of clusters desired, $k$, and the termination condition, $\tau$, initialize the means, update the clusters, and refine the estimates until converged. To ensure that the process terminates, stop after 100 iterations. 

When complete, return an array of distances from the current means to the previous means as well as the last vector of means.

<!-- grade_id: cell-92ac307eaadf9b28, points: 0, grade: false, solution: true -->
```ruby
def k_means(data, k, tau = 0.001)
  dists = []
  last_means = []
  curr_means = []
  iterations = 100
  last_means = init_cluster data, k
  
  iterations.times do |i|
    pre_means = last_means
    z = assign_cluster data, last_means
    curr_means = calculate_means z, data
    last_means = curr_means
    dist = cluster_dist_1 pre_means, curr_means
    
    dists << dist
    break if dist <= tau
   end
  
  return [dists, last_means]
end


def cluster_dist_1(m0, m1)
  dist = 0.0
  n = m0.size

  n.times do |i|
    dist_temp =  calc_dist_2 m0[i], m1[i]
    dist += dist_temp / n
  end
  dist
end

def calc_dist_2(point1, point2)
  dist = (point2["x1"] - point1["x1"]) ** 2 + (point2["x2"] - point1["x2"]) ** 2
  return dist
end

```

Output:
```
:calc_dist_2
```

<!-- grade_id: cell-4e1d634d3eee050f, points: 10, grade: true, solution: false -->
```ruby
def test_24_1()
  dataset = create_cluster_dataset()    
  dists, means, z = k_means dataset, 3, 0.001
  
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  df = Daru::DataFrame.new({iters: iters, dists: dists})
  df.plot(type: :line, x: :iters, y: :dists) do |plot, diagram|
    plot.x_label "X"
    plot.y_label "Mean Dist"
    diagram.title "Cluster Convergence"
    plot.legend false
  end
end

test_24_1()
```

Output:
```
#<Nyaplot::Plot:0x00000000014c9a40 @properties={:diagrams=>[#<Nyaplot::Diagram:0x000000000136a488 @properties={:type=>:line, :options=>{:x=>:iters, :y=>:dists, :title=>"Cluster Convergence"}, :data=>"ef846682-26c5-40ef-a02e-3b902b0318d5"}, @xrange=[0, 11], @yrange=[0.0, 12.285779548346301]>], :options=>{:x_label=>"X", :y_label=>"Mean Dist", :legend=>false, :zoom=>true, :width=>700, :xrange=>[0, 11], :yrange=>[0.0, 12.285779548346301]}}>
```

## Question 2.5 (15 points)
Plot the clusters that result for $k = 3, 5, 6, 10, 20$ clusters.

<!-- grade_id: cell-cc692e56ba45c1d7, points: 3, grade: true, solution: false -->
```ruby
def test_25_1(k)
  dataset = create_cluster_dataset()    
  dists, means, z = k_means dataset, k, 0.001
  
  assert_equal(k, means.size)
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  plot_clusters(dataset)
end

test_25_1(3)
```

Output:
```
#<Nyaplot::Plot:0x0000000000fac058 @properties={:diagrams=>[#<Nyaplot::Diagram:0x000000000220bca8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(27,158,119)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
```

<!-- grade_id: cell-6ed65b7a4cf777eb, points: 3, grade: true, solution: false -->
```ruby
test_25_1(5)
```

Output:
```
#<Nyaplot::Plot:0x000000000228a1c0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x000000000223b908 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>3, :color=>"rgb(251,180,174)", :tooltip_contents=>[3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3]}, :data=>"0337e0e6-82cc-4500-bf3d-c6326a7181f2"}, @xrange=[-6.029916383790382, -1.5013676292788314], @yrange=[1.2227259500074688, 7.531132609037203]>, #<Nyaplot::Diagram:0x00000000021ded48 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(179,205,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 
```

<!-- grade_id: cell-a28d8170d846f35c, points: 3, grade: true, solution: false -->
```ruby
test_25_1(6)
```

Output:
```
#<Nyaplot::Plot:0x000000000227d858 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000022584e0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(141,211,199)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"efdc0d95-9642-42c7-8b96-e50871c49492"}, @xrange=[-5.969358117731153, -1.6830537572213795], @yrange=[1.0952856617195168, 3.6692205529406197]>, #<Nyaplot::Diagram:0x0000000002222ae8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>3, :color=>"rgb(255,255,179)", :tooltip_contents=>[3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3]}, :data=>"2bae6099-5b83-404b-bf4a-60d180ba8f19"}, @xrange=[-5.802880395559711, -2.3197693631233984], @yrange=[3.6682364400358054, 6.5436326819646125]>, #<Nyaplot::Diagram:0x
```

<!-- grade_id: cell-6b0aa8259e513e3a, points: 3, grade: true, solution: false -->
```ruby
test_25_1(10)
```

Output:
```
#<Nyaplot::Plot:0x00000000022c2ac0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000022990d0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>3, :color=>"rgb(251,180,174)", :tooltip_contents=>[3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3]}, :data=>"f35f0081-e1ce-408c-8801-1e3d88f008a6"}, @xrange=[-4.988977782895842, -3.339054804294679], @yrange=[4.082737979694362, 8.079123643962838]>, #<Nyaplot::Diagram:0x000000000227ad60 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(179,205,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"8fac5eaf-b9fc-4fcd-a394-2d5544024d80"}, @xrange=[-6.36520132992231, -3.36474655690511], @yrange=[1.7778040098654135, 4.340824860849644]>, #<Nyaplot::Diagram:0x0000000002258508 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(204,235,197)", :to
```

<!-- grade_id: cell-686a77c7a91acdb9, points: 3, grade: true, solution: false -->
```ruby
test_25_1(20)
```

Output:
```
#<Nyaplot::Plot:0x0000000002115330 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000020b1df8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>3, :color=>"rgb(102,194,165)", :tooltip_contents=>[3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3, 3]}, :data=>"fb738738-70fb-4171-baf5-2213fc0ebaee"}, @xrange=[-5.899137162231485, -3.735329688266986], @yrange=[3.9714812642379194, 6.396808045305319]>, #<Nyaplot::Diagram:0x000000000200c448 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(252,141,98)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"ced38caa-662d-4153-8363-8a31a97a0efc"}, @xrange=[-6.5571464299445275, -3.951161412251332], @yrange=[2.5218145902011297, 4.218409190412217]>, #<Nyaplot::Diagram:0x0000000001fae2a8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>13, :color=>"rgb(141,160,203)", :tooltip_contents=>[13, 13, 13, 
```

## Question 2.6 (5 points)

Answer some questions about clustering. 

1. Which of cluster values fit the data best (visually)
 * (A) $k = 3$
 * (B) $k = 5$
 * (C) $k = 6$
 * (D) $k = 10$
 * (D) $k = 20$ 
1. When you set $k$ to a higher value than the ideal number of clusters what happens?
 * (A) Algorithm terminates because ideal number of clusters is a required argument
 * (B) All clusters except the ideal number are set to zero
 * (C) Clusters are split into smaller pieces
 * (D) Clusters are regularized out
1. Go back and try running the clusterings again. You may see that the obtained clusters are different in a few runs. Why? Is $k$-means is sensitive to the initialization?
 * (A) No, it is because the data is being randomly re-generated each time
 * (B) Yes, but only when $k$ is more than the ideal number of clusters
 * (C) No, $k$-means is a convex optimization problem and therefore insensitive to starting point
 * (D) Yes, it is sensitive to the initially guessed means
1. As discussed in the lecture, which of the following captures the biggest difference between $k$-means and the EM-algorithm for gaussian mixture models (with unit variance)?
 * (A) In the EM-algorithm, the $z$ vector is a probability and not just 0 or 1
 * (B) In the EM-algorithm, the means are selected from a normal distribution rather than a uniform distribution
 * (C) The K-means algorithm has no E step or M step, but EM has both 
 * (D) No difference, the algorithms are the same
1. Which datasets are well-suited to $k$-means clustering?
 * (A) All datasets, $k$-means is a low-bias algorithm
 * (B) High-dimensional, sparse vectors with multi-modal normal distributions
 * (C) Low-dimensional, dense vectors with multi-modal normal distributions
 * (D) Only toy datasets like those in this assignment
 
 
If your answers were all A, A, A, A, A, then write the following:

```ruby
def answer_51()
  answers = ["A", "A", "A", "A", "A"]
  return answers
end
```

<!-- grade_id: cell-90a0ff1dc7e47bb7, points: 0, grade: false, solution: true -->
```ruby
def answer_26()
  # BEGIN YOUR CODE
#   raise NotImplementedError.new()
  #END YOUR CODE
  answers = ["B", "C", "B", "A", "B"]
  return answers
end
```

Output:
```
:answer_26
```

<!-- grade_id: cell-769dbf6d94ea8044, points: 5, grade: true, solution: false -->
```ruby
t_answers = answer_26()

assert_not_nil t_answers, "1"
assert_true(t_answers.is_a?(Array))
assert_equal(5, t_answers.size)
assert_true(t_answers.all? {|a| a.size == 1 and a =~ /[A-Z]/})

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```