# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 6: Kernels


```ruby
require './assignment_lib'
```

Output:
```
"if(window['d3'] === undefined ||\n   window['Nyaplot'] === undefined){\n    var path = {\"d3\":\"https://cdnjs.cloudflare.com/ajax/libs/d3/3.5.5/d3.min\",\"downloadable\":\"http://cdn.rawgit.com/domitry/d3-downloadable/master/d3-downloadable\"};\n\n\n\n    var shim = {\"d3\":{\"exports\":\"d3\"},\"downloadable\":{\"exports\":\"downloadable\"}};\n\n    require.config({paths: path, shim:shim});\n\n\nrequire(['d3'], function(d3){window['d3']=d3;console.log('finished loading d3');require(['downloadable'], function(downloadable){window['downloadable']=downloadable;console.log('finished loading downloadable');\n\n\tvar script = d3.select(\"head\")\n\t    .append(\"script\")\n\t    .attr(\"src\", \"http://cdn.rawgit.com/domitry/Nyaplotjs/master/release/nyaplot.js\")\n\t    .attr(\"async\", true);\n\n\tscript[0][0].onload = script[0][0].onreadystatechange = function(){\n\n\n\t    var event = document.createEvent(\"HTMLEvents\");\n\t    event.initEvent(\"load_nyaplot\",false,false);\n\t    win
```

## Question 1.1: Prepare Dataset (1 points)

We will generate and plot a synthetic dataset to illustrate Kernel learning. Here is a pre-defined dataset for us to use in this exercise. Paste **your** dot and norm function in the cell below.

```ruby
def plot_spiral_dataset()
  spiral = spiral_dataset()["data"]
  x1 = spiral.map {|r| r["features"]["x1"]}
  x2 = spiral.map {|r| r["features"]["x2"]}
  target = spiral.map {|r| r["label"]}
  target.size  
  df = Daru::DataFrame.new({x1: x1, x2: x2, target: target})
  df.to_category :target
  df.plot(type: :scatter, x: :x1, y: :x2, categorized: {by: :target, method: :color}) do |plot, diagram|
    plot.xrange [-8,8]
    plot.x_label "X1"
    plot.yrange [-8,8]  
    plot.y_label "X2"
    plot.legend false
  end
end

plot_spiral_dataset()
```

Output:
```
#<Nyaplot::Plot:0x00000000031a9e08 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000030f50c0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1.0, :color=>"rgb(251,180,174)", :tooltip_contents=>[1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0]}, :data=>"54e746f8-6b91-48b4-b221-0432071e1633"}, @xrange=[-5.0, 6.0], @yrange=[-5.5, 6.5]>, #<Nyaplot::Diagram:0x00000000035b8da0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0.0, :color=>"rgb(179,205,227)", :tooltip_contents=>[0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0
```

<!-- grade_id: cell-ff4b2c6b485a3e03, points: 0, grade: false, solution: true -->
```ruby
# BEGIN YOUR CODE
def dot x, w
  res = 0.0
  x.each do |k, v|
    if w[k] != nil
      res += v * w[k]
    end
  end
  return res  
end

def norm w
  return Math.sqrt(dot(w,w))
end
#END YOUR CODE
```

Output:
```
:norm
```

<!-- grade_id: cell-aabc301bebd79fdd, points: 1, grade: true, solution: false -->
```ruby
def test_11()
  assert_in_delta 2.0, norm({"a" => 1.41421, "b" => 1.41421}), 1e-2
  assert_in_delta 2.0, norm({"a" => -1.41421, "b" => 1.41421}), 1e-2
  assert_in_delta 0.0, norm({}), 1e-2

  assert_in_delta 6.0, dot({"a" => 2.0}, {"a" => 3.0}), 1e-6
  assert_in_delta 6.0, dot({"a" => 2.0}, {"a" => 3.0, "b" => 4.0}), 1e-6
  assert_equal 0.0, dot({}, {})
  assert_equal 0.0, dot({"a" => 1.0}, {"b" => 1.0})
end

test_11()
```

## Question 1.2 (1 Point)
Copy **YOUR** implementation of ```StochasticGradientDescent``` from [Assignment 5](../assignment-5/assignment-5.ipynb) into the following cell.

<!-- grade_id: cell-512b41c089c45d6b, points: 0, grade: false, solution: true -->
```ruby
# BEGIN YOUR CODE
# raise NotImplementedError.new()

class StochasticGradientDescent
  attr_reader :weights
  attr_reader :objective
  def initialize obj, w_0, lr = 0.01
    @objective = obj
    @weights = w_0
    @n = 1.0
    @lr = lr
  end
  def update x
    # BEGIN YOUR CODE
    curr_lr = @lr / Math.sqrt(@n)
    @objective.func(x, @weights)
    grad = @objective.grad(x, @weights)
    @weights = update_weights(@weights, grad, curr_lr)
    @n += 1
    #END YOUR CODE
  end
  
  def update_weights(w, dw, lr)
    w_copy = w.clone()
    dw_copy = dw.clone()
  
    dw_copy.each do |k, v|
      dw_copy[k] *= lr
    end
  
    w_copy.each do |k, v|
      if dw_copy.key?(k)
        w_copy[k] -= dw_copy[k]
      end
    end
    w_copy
  end
  
end
#END YOUR CODE
```

Output:
```
:update_weights
```

<!-- grade_id: cell-decbf7a78199df43, points: 1, grade: true, solution: false -->
```ruby
### Hidden Test (See Test 1.1 from Assignment 5) ###
assert_not_nil(StochasticGradientDescent.class)
```

## Question 1.3 (5 Points)


Refactor **your** implementation of the AUC metric to return the false positive, true positive values for all scores in addition to the AUC. We will use these to plot the ROC curve.

<!-- grade_id: cell-b97ece98e0d680af, points: 0, grade: false, solution: true -->
```ruby
def roc_curve(scores)
  # BEGIN YOUR CODE
  res = Array.new
  
  min = (scores.min{|a,b| a[0] <=> b[0]})[0]
  max = (scores.max{|a,b| a[0] <=> b[0]})[0]
  num = 1000.0
  tp = Array.new
  fp = Array.new
  step = (max - min) / num
  
  (min..max).step(step).to_a.each do |threshold|
    false_neg = 0.0 
    true_pos = 0.0
    true_neg = 0.0
    false_pos = 0.0
    scores.each do |record|
      if(record[1] == 0 or record[1] == -1) and record[0] >= threshold
        false_pos += 1.0
      end
      if(record[1] == 0 or record[1] == -1) and record[0] < threshold
        true_neg += 1.0
      end
      if(record[1] == 1 and record[0] >= threshold)
        true_pos += 1.0
      end
      if(record[1] == 1 and record[0] < threshold)
        false_neg += 1.0
      end
    end
    fp_temp = false_pos / (false_pos + true_neg)
    tp_temp = true_pos / (true_pos + false_neg)
    fp << fp_temp
    tp << tp_temp
    
    res << [fp_temp, tp_temp]    
  end
  auc = 0.0
  sorted = res.sort{|a,b| b[0] <=> a[0]}
  (1..sorted.length - 1).each do |index|
    width = sorted[index - 1][0] - sorted[index][0]
    height = sorted[index][1]
    auc += width * height
  end

  #END YOUR CODE
  return [fp, tp, auc]
end
```

Output:
```
:roc_curve
```

<!-- grade_id: cell-bce19441b35b0abd, points: 5, grade: true, solution: false -->
```ruby
def test_13()
  good_model = [[0.9, 1], [0.89, 1], [0.7, 0], [0.8, 1], [0.8, 0], [0.7, 1], [0.6, 0], [0.5, 0], [0.1, 0]]
  assert_true(roc_curve(good_model).last > 0.8)
  assert_true(roc_curve(good_model).last < 1)
  
  srand(777)
  ok_model = Array.new(100) {|i| [100 - i, (rand < (100 - i) / 100.0) ? 1 : 0] }
  ok_auc = roc_curve(ok_model).last
  assert_in_delta(0.8631239935587761, ok_auc, 1e-3)
  
  bad_model = Array.new(1000) {|i| [1000 - i, rand < 0.5 ? 1 : 0] }
  bad_auc = roc_curve(bad_model).last
  assert_in_delta(0.5, bad_auc, 5e-2)
  
  fp, tp, auc = roc_curve(good_model)
  
  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
    diagram.title("Good Model (AUC > 0.8)")
    plot.legend(true)
  end
  
  fp, tp, auc = roc_curve(bad_model)
  plot.add(:line, fp, tp).color("red").title("Bad Model (AUC = 0.5)")
  
  plot.show()
end

test_13()
```

Output:
```
#<Nyaplot::Frame:0x0000000002ff70b0 @properties={:panes=>[#<Nyaplot::Plot:0x0000000002931a50 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003578e30 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y, :title=>"Good Model (AUC > 0.8)"}, :data=>"ac0ae7ae-6436-4d13-8b10-9ae4abd4c791"}, @xrange=[0.0, 1.0], @yrange=[0.25, 1.0]>, #<Nyaplot::Diagram:0x0000000002ff7bf0 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>"red", :title=>"Bad Model (AUC = 0.5)"}, :data=>"f661c441-8af2-4325-8543-52b3323d7372"}, @xrange=[0.0019723865877712033, 1.0], @yrange=[0.0, 1.0]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :legend=>true, :zoom=>true, :width=>800, :xrange=>[0.0, 1.0], :yrange=>[0.0, 1.0]}}>], :data=>{"ac0ae7ae-6436-4d13-8b10-9ae4abd4c791"=>#<Nyaplot::DataFrame:0x00000000035b11e0 @name="ac0ae7ae-6436-4d13-8b10-9ae4abd4c791", @rows=[{:x=>1.0, :y=>1.0}, {:x=>0.8, :y=>1.0}, {:x=>0.8, :y=>1.0}, {:x=>0.8, :y=>1.0}, {:x=>0.8, :y=>1.0}, {:
```

##  Question 2.1 (10 Points)

Implement the Hinge Loss regularized linear model. The prediction function is the same as linear regression

<!-- grade_id: cell-2f4f3a643f0b678c, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def initialize reg
    @reg_param = reg
  end
  
  def predict row, w
    # BEGIN YOUR CODE
    res = 0.0
    row["features"].each do |key, value|
      if(w.key? (key))
        res += value * w[key]
      end
    end
    res
    #END YOUR CODE
  end
  
  def adjust w
  end
end
```

Output:
```
:adjust
```

<!-- grade_id: cell-67fe4fb894a07ef0, points: 10, grade: true, solution: false -->
```ruby
## Tests predict function
def test_21()
  t_loss = HingeLossClassifier.new 0.0
  assert_equal(1.0, t_loss.predict({"features" => {"x1" => 1.0}}, {"x1" => 1.0}))
  assert_equal(-1.0, t_loss.predict({"features" => {"x1" => -1.0}}, {"x1" => 1.0}))
  assert_equal(-14.0, t_loss.predict({"features" => {"x1" => 2.0}}, {"x1" => -7.0}))
end

test_21()
```

## Question 2.2 (10 Points)

Implement the objective function, which is defined as follows for a mini-batch of $n$ examples:

# $L(w,X) = \frac{\lambda}{2} \left\lVert w \right\rVert ^ 2 + \frac{1}{n} \sum_{i} \max \left[ 0, 1 - y_i \widehat{y}_i\right]$

where ```@reg_param``` is the value $\lambda$.

Remember that for hinge loss, the label must be $y = \pm 1$.

<!-- grade_id: cell-92a64a66671de794, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def func data, w
    # BEGIN YOUR CODE
#     raise NotImplementedError.new()
    res = 0.0
    size = data.length
    data.each do |record|
      label = record["label"] == 0 ? -1 : record["label"]
      predict_val = predict record, w
      update_val = ([0, 1 - label * predict_val].max)
      res += update_val
    end
    
    res /= size
    res += (@reg_param * norm(w) ** 2 / 2)
    #END YOUR CODE
    res
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-2dbf2caa8b69e62c, points: 10, grade: true, solution: false -->
```ruby
def test_22()
  ## Tests objective function
  t_loss = HingeLossClassifier.new 0.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_equal(0.0, t_loss.func(t_data, {"x1" => 1.0}), "1")
  assert_equal(2.0, t_loss.func(t_data, {"x1" => -1.0}), "2")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 0.0}), "3")

  t_data += [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 1.0}), "4")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => -1.0}), "5")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 0.0}), "6")

  t_loss = HingeLossClassifier.new 27.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_in_delta(675.0, t_loss.func(t_data, {"x1" => 1.0, "x2" => 7.0}), 1e-2, "7")
  assert_in_delta(677.0, t_loss.func(t_data, {"x1" => -1.0, "x2" => -7.0}), 1e-2, "8")
end
test_22()
```

## Question 2.3 (10 Points)

Implement the Hinge Loss classifier gradient method.


<!-- grade_id: cell-31495bf0a5d226c2, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def grad data, w
    g = Hash.new {|h,k| h[k] = 0.0}
    # BEGIN YOUR CODE
    size = data.length
    data.each do |record|
      label = record["label"] == 0 ? -1 : record["label"]
      predict_val = predict record, w
      record["features"].each do |key, value|
        if(!g.key?(key))
          g[key] = 0.0
        end
        if ((label * predict_val) < 1)
          g[key] -= label * value / size
        end
      end
    end
    
    g.each do |key, value|
      g[key] = value + @reg_param * w[key]
    end
    
    #END YOUR CODE
    return g
  end
end
```

Output:
```
:grad
```

<!-- grade_id: cell-ce1477dcf71b4fbe, points: 10, grade: true, solution: false -->
```ruby
def test_33()
  t_loss = HingeLossClassifier.new 0.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_equal({"x1" => 0.0}, t_loss.grad(t_data, {"x1" => 1.0}), "1")
  assert_equal({"x1" => -1.0}, t_loss.grad(t_data, {"x1" => -1.0}), "2")
  assert_equal({"x1" => -1.0}, t_loss.grad(t_data, {"x1" => 0.0}), "3")

  t_data += [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_equal({"x1" => 0.5}, t_loss.grad(t_data, {"x1" => 1.0}), "4")
  assert_equal({"x1" => -0.5}, t_loss.grad(t_data, {"x1" => -1.0}), "5")
  assert_equal({"x1" => 0.0}, t_loss.grad(t_data, {"x1" => 0.0}), "6")

  t_loss = HingeLossClassifier.new 11.7
  t_data = [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_in_delta(12.7, t_loss.grad(t_data, {"x1" => 1.0})["x1"], 1e-2, "7")
  assert_in_delta(-11.7, t_loss.grad(t_data, {"x1" => -1.0})["x1"], 1e-2, "8")
  assert_in_delta(1.0, t_loss.grad(t_data, {"x1" => 0.0})["x1"], 1e-2, "9")

end
test_33()
```

## Question 2.4 (10 Points)

Train and test the Hinge Loss classifier on the same (training) sprial dataset. Return an array or (score, label) pairs. Try different learning rates, epochs, batch sizes, and regularization parameters. 

<!-- grade_id: cell-e4ab95acd13266d2, points: 0, grade: false, solution: true -->
```ruby
def train_linear_hinge(dataset)
  lr = HingeLossClassifier.new 0.5
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  
  sgd = StochasticGradientDescent.new lr, w, 0.05
  pairs_arr = Array.new
  
  # BEGIN YOUR CODE
  data = dataset["data"]
  batch_size = 100
  epoch_num = 50
  iter_per_epoch = data.length / batch_size
  iter_num = iter_per_epoch * epoch_num
  
  iter_num.times do 
    batch_data = data.sample(batch_size)
    sgd.update(batch_data)
    w = sgd.weights
  end
  
  data.each do |record|
    scores = lr.predict(record, w)
    pairs_arr << [scores, record["label"]]
  end
  
  
  pairs_arr
  
  #END YOUR CODE
end
```

Output:
```
:train_linear_hinge
```

<!-- grade_id: cell-57a39fde4daf78e8, points: 10, grade: true, solution: false -->
```ruby
def test_24()
  spiral = spiral_dataset()
  scores = train_linear_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.5, "AUC > 0.5")
  assert_true(auc > 0.52, "AUC > 0.52")  
  assert_true(auc < 0.7, "AUC < 0.7")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_24()
```

Output:
```
0.5442661281751514
#<Nyaplot::Frame:0x00000000035bfc90 @properties={:panes=>[#<Nyaplot::Plot:0x0000000003648a90 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000035e50a8 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y}, :data=>"8a36a2cb-14dd-46a8-b30c-53365358c567"}, @xrange=[0.0, 1.0], @yrange=[0.010309278350515464, 1.0]>, #<Nyaplot::Diagram:0x00000000035e41f8 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>:gray}, :data=>"291d1167-9a22-41b4-9407-286588f2a612"}, @xrange=[0, 1], @yrange=[0, 1]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :zoom=>true, :width=>700, :xrange=>[0.0, 1.0], :yrange=>[0, 1.0]}}>], :data=>{"8a36a2cb-14dd-46a8-b30c-53365358c567"=>#<Nyaplot::DataFrame:0x00000000035e6ef8 @name="8a36a2cb-14dd-46a8-b30c-53365358c567", @rows=[{:x=>1.0, :y=>1.0}, {:x=>0.9690721649484536, :y=>1.0}, {:x=>0.9587628865979382, :y=>1.0}, {:x=>0.9484536082474226, :y=>1.0}, {:x=>0.9381443298969072, :y=>1.0}, {:x=>0.9278350515
```

## Question 3.1 (10 Points)

Implement the linear kernel where $x_i$ and $x_j$ are rows in a dataset.

<!-- grade_id: cell-fcf79802d7f8c72d, points: 0, grade: false, solution: true -->
```ruby
class LinearKernel
  def func x_i, x_j
    # BEGIN YOUR CODE
    res = 0.0
    
    features_a = x_i["features"]
    features_b = x_j["features"]
    
    features_a.each do |key, value|
      if features_b.key?(key)
        res += value * features_b[key]
      end
    end
    
    res
    #END YOUR CODE
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-3342f2f3fed428da, points: 5, grade: true, solution: false -->
```ruby
def test_31()
  linear = LinearKernel.new
  
  kij = linear.func({"features" => {"a" => 2, "b" => 3}}, {"features" => {"b" => 7, "c" => 1}})  
  assert_in_delta 21, kij, 1e-2, "1"
  
  kij = linear.func({"features" => {"a" => 2, "b" => 3}}, {"features" => {"a" => -5, "b" => 2}})  
  assert_in_delta -4, kij, 1e-2, "1"
end

test_31()
```

## Question 3.2 (10 Points)

Implement the Gaussian Kernel, taking a bandwidth parameter $\sigma$. 

<!-- grade_id: cell-bd59856809fb8bc2, points: 0, grade: false, solution: true -->
```ruby
class GaussianKernel
  def initialize width
    @width = width
  end
  def func x_i, x_j
    # BEGIN YOUR CODE
#     res = 0.0
#     sec_res = 0.0
#     features_1 = x_i["features"]
#     features_2 = x_j["features"]
    
#     features_1.each do |key, value|
#       if features_2.key?(key)
#         update_val = (value - features_2[key]) ** 2
#         sec_res -= update_val * @width
#       end
#     end
    
#     res = Math.exp(sec_res)
#     res
    diff = Hash.new {|h,k| h[k] = 0.0}
    x_i["features"].each do |key, value|
      if x_j["features"][key] != nil
        diff[key] = value - x_j["features"][key]
      end
    end
    return Math.exp(-@width * (norm(diff)**2))
    #END YOUR CODE
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-ec6c13625f9e6801, points: 10, grade: true, solution: false -->
```ruby
def test_32()
  spiral = spiral_dataset()
  kernel = GaussianKernel.new 1.0
  
  test_v = kernel.func spiral["data"][0], spiral["data"][2]
  assert_true(test_v > 0.19, "1")
  assert_true(test_v < 0.25, "2")
end

test_32()
```

## Question 3.3 (10 Points): Gram Matrix

Implement the Gram Matrix as a preprocessing step which is first trained and then used to transform a dataset. Given a training set $x_i \in X$ and a kernel function $k$ the Gram Matrix is defined as:

# $G_{i,j} = k(x_i, x_j)$

For any new datapoint z, the _Gram vector_ is defined as follows:

# $g(z) = \left( k(z, x_i), \dots, k(z, x_n) \right)$

Note that on the training $g(x_i) = G_{i,\cdot}$. 

1. Implement a ```train``` method that caches data for evaluation. 
1. Implement the ```func``` method that takes a single row and the gram vector as a row 

<!-- grade_id: cell-6dec8ec0ea6c2f66, points: 0, grade: false, solution: true -->
```ruby
class GramMatrix
  def initialize kernel
    @kernel = kernel
    @data = Array.new
  end
  
  def train data
    # BEGIN YOUR CODE
    @data = data
    # END YOUR CODE
  end
  
  def transform z
    # BEGIN YOUR CODE
    res = Hash.new
    if z["label"] != nil
      res["label"] = z["label"]
    end
    res["features"] = Hash.new
    
    size = @data.length
    size.times do |i|
      key = "k_#{i}"
      res["features"][key] = @kernel.func(z, @data[i])
    end
    res
    # END YOUR CODE
  end
end
```

Output:
```
:transform
```

<!-- grade_id: cell-f9dc67d1e3802453, points: 10, grade: true, solution: false -->
```ruby
def test_33()
  lin = LinearKernel.new
  data =  [
    {"features" => {"a" => 2, "b" => 3}, "label" => 1}, 
    {"features" => {"b" => 7, "c" => 1}, "label" => 7},
    {"features" => {"a" => 2, "b" => 3}, "label" => 2}
  ]
  
  z = {"features" => {"a" => -5, "b" => 2}} 

  gmat = GramMatrix.new lin
  gmat.train data

  k_x1 = gmat.transform data[1]
  assert_equal 3, k_x1["features"].size, "1"
  assert_equal 7, k_x1["label"], "2"
  assert_equal dot(data[1]["features"], data[1]["features"]), k_x1["features"]["k_1"], "3"
  assert_equal 21, k_x1["features"]["k_0"], "4"
  k_x2 = gmat.transform z
  assert_equal 3, k_x1["features"].size, "5"
  assert_in_delta -4, k_x2["features"]["k_0"], "6"
end

test_33()
```

## Question 4.1 (10 Points)

Test the Hinge Loss with Linear Kernel on the spiral dataset, return the false positive, true positive, and auc values. Do you see any improvements with this kernelized version? 

<!-- grade_id: cell-6f585f8a01072b23, points: 1, grade: true, solution: true -->
```ruby
def train_linear_kernel_hinge(dataset)
  lr = HingeLossClassifier.new 0.1
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  pairs_arr = Array.new
  
  # BEGIN YOUR CODE
  linear_kernel = LinearKernel.new
  gram_matrix = GramMatrix.new linear_kernel
  data = dataset["data"]
  data_size = data.length
  gram_matrix.train(data)
  transformed = Array.new
  
  data_size.times do |i|
    transformed[i] = gram_matrix.transform(data[i])
  end
  
  sgd = StochasticGradientDescent.new lr, w, 0.05
  batch_size = 50
  iter_per_epoch = transformed.length / batch_size
  epoch_num = 100
  iter_num = epoch_num * iter_per_epoch
  
  iter_num.times do |i|
    batch_data = transformed.sample(batch_size)
    sgd.update(batch_data)
    weight = sgd.weights
  end
  
  transformed.each do |record|
    score = lr.predict(record, w)
    label = record["label"]
    pairs_arr << [score, label]
  end
  pairs_arr
  #END YOUR CODE
end
```

Output:
```
:train_linear_kernel_hinge
```

<!-- grade_id: cell-444aa68056e01128, points: 10, grade: true, solution: false -->
```ruby
def test_41()
  spiral = spiral_dataset()
  scores = train_linear_kernel_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.45, "AUC > 0.45")
  assert_true(auc < 0.7, "AUC < 0.7")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_41()
```

Output:
```
0.5010096715910297
#<Nyaplot::Frame:0x0000000003ce3c60 @properties={:panes=>[#<Nyaplot::Plot:0x0000000003c15f90 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003ce5c90 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y}, :data=>"6bd2a2a1-2b96-4f22-bbc5-bef2cf037f7c"}, @xrange=[0.010309278350515464, 1.0], @yrange=[0.0, 1.0]>, #<Nyaplot::Diagram:0x0000000003ce4430 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>:gray}, :data=>"7aeb917c-792a-4163-99e4-072b22129958"}, @xrange=[0, 1], @yrange=[0, 1]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :zoom=>true, :width=>700, :xrange=>[0, 1.0], :yrange=>[0.0, 1.0]}}>], :data=>{"6bd2a2a1-2b96-4f22-bbc5-bef2cf037f7c"=>#<Nyaplot::DataFrame:0x000000000247ec10 @name="6bd2a2a1-2b96-4f22-bbc5-bef2cf037f7c", @rows=[{:x=>1.0, :y=>1.0}, {:x=>1.0, :y=>0.9896907216494846}, {:x=>1.0, :y=>0.9896907216494846}, {:x=>1.0, :y=>0.9896907216494846}, {:x=>1.0, :y=>0.9896907216494846}, {:x=>1.0, :y=>0.9
```

## Question 3.3 (5 points)

Is the AUC for Linear Kernel much better than before? Why or Why not?

<!-- grade_id: cell-6cf57bec611f2fa2, points: 0, grade: false, solution: true -->
```ruby
def train_gaussian_kernel_hinge(dataset)
  lr = HingeLossClassifier.new 0.1
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  pairs_arr = Array.new
  
  # BEGIN YOUR CODE
  kernel = GaussianKernel.new 2.0
  gram_matrix = GramMatrix.new kernel
  data = dataset["data"]
  data_size = data.length
  gram_matrix.train(data)
  transformed = Array.new
  
  data_size.times do |i|
    transformed[i] = gram_matrix.transform(data[i])
  end
  
  sgd = StochasticGradientDescent.new lr, w, 0.3
  
  batch_size = 50
  num_epoch = 100
  iter_per_epoch = transformed.length / batch_size
  iter_num = num_epoch * iter_per_epoch
  iter_num.times do |i|
    batch_data = transformed.sample(batch_size)
    sgd.update(batch_data)
    weight = sgd.weights
  end
  
  transformed.each do |record|
    score = lr.predict(record, w)
    label = record["label"]
    pairs_arr << [score, label]
  end
  pairs_arr
  
  #END YOUR CODE
end
```

Output:
```
:train_gaussian_kernel_hinge
```

<!-- grade_id: cell-e94797f3c52a8594, points: 10, grade: true, solution: false -->
```ruby
def test_42()
  spiral = spiral_dataset()
  scores = train_gaussian_kernel_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.5, "AUC > 0.5")
  assert_true(auc > 0.6, "AUC > 0.6")  
  assert_true(auc > 0.7, "AUC > 0.7")    
  assert_true(auc > 0.8, "AUC > 0.8")    
  assert_true(auc < 1.0, "AUC < 1.0")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_42()
```

Output:
```
0.5323626315230096
/usr/local/rvm/gems/ruby-2.5.1/bin/iruby:23:in `<main>'
/usr/local/rvm/gems/ruby-2.5.1/bin/ruby_executable_hooks:15:in `eval'
/usr/local/rvm/gems/ruby-2.5.1/bin/ruby_executable_hooks:15:in `<main>'
```

## Question 5.1 (7 Points)

Compare the un-kernelized version, linear kernel, and Gaussian kernels on this dataset. Provide answers to the following questions as an array below:

1. Which model had the best performance as measured by AUC
 * (A) Linear Kernel
 * (B) Gaussian Kernel
 * (C) Non-Kernel
 * (D) They were the same
1. If there are $N$ examples in the training set, what is the size of the Gram matrix during training?
 * (A) $O(N)$
 * (B) $O(N \log N)$
 * (C) $O(1)$
 * (D) $O(N^2)$
1. For linear classifier, if there are $N$ examples with $d$ features in the training set, what is the time complexity of scoring one example?
 * (A) $O(d)$
 * (B) $O(N^2)$
 * (C) $O(Nd)$
 * (D) $O(N \log d)$
1. For kernelized classifier, if there are $N$ examples with $d$ features in the training set, what is the time complexity of scoring one example?
 * (A) $O(N \log d)$
 * (B) $O(d)$
 * (C) $O(N^2)$
 * (D) $O(Nd)$
1. Why do you think the Gaussian kernel would have better performance for the spiral dataset?
 * (A) Hinge loss causes the linear kernel to overfit, but the Gaussian kernel is very smooth and therefore less prone to overfitting.
 * (B) Gaussian kernel did not perform better than anything else.
 * (C) The spiral dataset is non-linear and the Gaussian kernel can learn non-linear boundaries.
 * (D) The data follows a Gaussian distribution therefore Gaussian kernels should fit the best. 
 
 
If your answers were all A, A, A, A, A, then write the following:

```ruby
def answer_51()
  answers = ["A", "A", "A", "A", "A"]
  return answers
end
```

<!-- grade_id: cell-0e8a77df1540314c, points: 0, grade: false, solution: true -->
```ruby
def answer_51()
  # BEGIN YOUR CODE
  answers = ["B", "D", "A", "B", "C"]
  #END YOUR CODE
  return answers
end
```

Output:
```
:answer_51
```

<!-- grade_id: cell-f64f4eb643d0d952, points: 7, grade: true, solution: false -->
```ruby
t51_answers = answer_51()

assert_not_nil t51_answers, "1"
assert_true(t51_answers.is_a?(Array))
assert_equal(5, t51_answers.size)
assert_true(t51_answers.any? {|a| a.size == 1 and a =~ /[A-Z]/})

```

```ruby

```