# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 0 - Getting Started

This will be a quick introduction to the environment and notebooks.

```ruby
require 'test/unit/assertions'
include Test::Unit::Assertions
```

## Question 1.1: (2 Points)

In the cell below, make the function below return a 1.

<!-- grade_id: cell-724f56610cc85545, points: 0, grade: false, solution: true -->
```ruby
def q_11
  # BEGIN YOUR CODE
  # raise NotImplementedError.new()
  1
  #END YOUR CODE
end
```

<!-- grade_id: cell-7dd749d31d93660a, points: 1, grade: true, solution: false -->
```ruby
assert_equal(1, q_11())
```

<!-- grade_id: cell-524bb3f7d35c2456, points: 1, grade: true, solution: false -->
```ruby
# The following is a hidden test. You will not be able to run it, but it will affect your grade.
```

```ruby

```