# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 1: Environment Setup

This assignment tests your access to the environment and ability to run through the basic steps.


```ruby
require 'test/unit/assertions'
include Test::Unit::Assertions
```

Output:
```
Object
```

## Question 1: 50 Points

Calculate the mean of an array of numbers.

<!-- grade_id: cell-182ea74825cac608, points: 0, grade: false, solution: true -->
```ruby
def mean x
  # BEGIN YOUR CODE
  sum = 0.0
  for item in x
    sum = sum + item
  end
  return sum / x.size
#   raise NotImplementedError.new()
  #END YOUR CODE
end
```

Output:
```
:mean
```

<!-- grade_id: cell-872ecd07d4058711, points: 25, grade: true, solution: false -->
```ruby
### TEST ###
test_1 = [3.0, 4.0, 5.0]
assert_equal(4.0, mean(test_1))
```

<!-- grade_id: cell-a29b65bc2c64efbf, points: 25, grade: true, solution: false -->
```ruby
### TEST ###
def mean_solution x
  x.inject(0.0) {|u,v| u += v} / x.size.to_f 
end
assert_equal(mean_solution(test_1), mean(test_1))
```

---
## Question 2: 50 Points

Calculate the entropy of a distribution of positive values. First, normalize them to sum to 1. Next, take special care to handle the case where a probability value is zero. Assume we use ruby's ```Math.log``` function.

The formula is as follows:

# $H(X) = -\sum_{x \in X} P(x) \log P(x)$


<!-- grade_id: cell-ad0650cc8b09869e, points: 0, grade: false, solution: true -->
```ruby
def entropy freq_x
  # BEGIN YOUR CODE
  sum = 0.0
  for item in freq_x
    sum += item
  end
  if sum == 0
    return 0.0
  end
  freq_x.collect!{|x| x / sum}
  entropy = 0.0
  for item in freq_x
    if item == 0.0
      next
    end
    entropy -= (item) * Math.log(item)
  end
  entropy
#   raise NotImplementedError.new()
  #END YOUR CODE
end
```

Output:
```
:entropy
```

<!-- grade_id: cell-dcc0a20e2049b7ea, points: 13, grade: true, solution: false -->
```ruby
### TEST ###
test_2_1 = [4, 4, 4, 4]
assert_in_delta(-Math.log(0.25), entropy(test_2_1), 1e-5)
```

<!-- grade_id: cell-dc11f755a63c6709, points: 13, grade: true, solution: false -->
```ruby
### TEST ###
test_2_2 = [4, 4, 0, 4]
assert_in_delta(-Math.log(1.0 / 3.0), entropy(test_2_2), 1e-5)
```

<!-- grade_id: cell-b6fa0c4bd9ffa9e6, points: 12, grade: true, solution: false -->
```ruby
### TEST ###
test_2_3 = [4]
assert_in_delta(0.0, entropy(test_2_3), 1e-5)
```

<!-- grade_id: cell-f3203b70d5afd577, points: 12, grade: true, solution: false -->
```ruby
### TEST ###
test_2_4 = [0,0,0,0,0]
assert_in_delta(0.0, entropy(test_2_4), 1e-5)
```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```

```ruby

```